--- 
title_meta  : Chapitre 4
title       : Facteurs
description : Très souvent, les données sont qualitatives. Par exemple, les individus sont soit des hommes, soit des femmes. Dans R, les données qualitatives sont stockées dans des facteurs. Compte-tenu de l'importance de ces facteurs dans l'analyse des données, vous allez commencer dès maintenant à apprendre à créer, extraire et comparer ces facteurs ! 
framework   : datacamp
mode        : selfcontained

--- type:NormalExercise xp:100 key:f5da7b110b
## Qu'est ce qu'un facteur et pourquoi devrais-je en utiliser ? 

Dans ce chapitre, vous plongez dans le monde magnifique des **facteurs**. 

Le terme de `factor` se réfère à un type de données statistiques utilisé pour stocker les variables nominales (on dit aussi qualitatives ou catégorielles). La différence entre une variable catégorielle et une variable continue, c'est qu'une variable catégorielle ne peut contenir qu'un **nombre limité de catégories**. Une variable continue, quant à elle, peut correspondre à un nombre infini de valeurs. 

Il est important que R sache s'il traite une variable continue ou une variable catégorielle, car les modèles statistiques, que vous allez développer dans l'avenir traitent les deux types différemment. (Vous verrez plus tard pourquoi c'est le cas.)

La variable 'Sexe' en est un bon exemple d'une variable qualitative. Un individu humain peut être 'Masculin' ou 'Féminin', faisant abstraction des intersexués. 'Masculin' et 'Féminin' sont, de manière simplifiée, les deux valeurs de la variable qualitative 'Genre' et chaque  valeur de cette variable ne peut contenir que 'Masculin' et 'Féminin'.

Pour créer des `factor` dans R, utilisez la fonction [`factor()`] (http://www.rdocumentation.org/packages/base/functions/factor). La première chose que vous devez faire est de créer un vecteur qui contient toutes les observations qui appartiennent à un nombre limité de catégories. Par exemple, `genre_vector` contient le sexe de 5 personnes différentes : 

```
genre_vector <- c("Masculin","Feminin","Feminin","Masculin","Masculin")
```

Il est clair ici qu'il y a 2 catégories, ou dans R des **'factor levels'**, il s' agit de "Masculin" et "Feminin".

La fonction [`factor()`] (http://www.rdocumentation.org/packages/base/functions/factor) va encoder le vecteur comme un factor : 

```
factor_genre_vector <- factor(genre_vector)
``` 

*** =instructions
Assigner, à `factor_genre_vector`, le vecteur de caractères `genre_vector` converti en un **factor**. Regardez la console et notez que R va afficher les niveaux de facteur en dessous des valeurs. Ils sont précédé de la mention **levels**.

*** =hint
Il suffit d'utiliser la fonction [`factor()`] (http://www.rdocumentation.org/packages/base/functions/factor) sur `genre_vector`. Jetez un oeil aux instructions, la réponse est déjà là, quelque part...

*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
genre_vector <- c("Masculin", "Feminin", "Feminin", "Masculin", "Masculin")

# Fabriquez factor_genre_vector en utilisant 'factor()'
factor_genre_vector <-
```

*** =solution
```{r eval=FALSE}
genre_vector <- c("Masculin", "Feminin", "Feminin", "Masculin", "Masculin")

# Fabriquez factor_genre_vector en utilisant 'factor()'
factor_genre_vector <- factor(genre_vector)
```

*** =sct
```{r eval=FALSE}
test_object("factor_genre_vector",
            undefined_msg = "Faites attention &#224; bien d&#233;finir la variable <code>factor_genre_vector</code>.",
            incorrect_msg = "Avez-vous bien pass&#233; <code>genre_vector</code> en tant que param&#232;tre de la fonction factor, pour stocker le r&#233;sultat dans <code>factor_genre_vector</code> ? ")
success_msg("Bien ! Si vous voulez en savoir plus sur la fonction <code>factor()</code>, n'h&#233;sitez pas &#224; taper <code> ? factor</code> dans la console. Cela ouvrira une page d'aide. On passe &#224; l'exercice suivant.");
```

*** =skills
1

--- type:NormalExercise xp:100 key:035c13837d
## Qu'est ce qu'un facteur et pourquoi devrais-je en utiliser ? (2)

Il existe deux types de variables nominales : les variables non ordonnées et les variables ordonnées. 

Une variable non ordonnée est une variable catégorielle sans ordre implicite. Cela signifie qu'il est impossible de dire qu' une valeur est supérieure à l'autre. Par exemple, pensez à la variable catégorielle `animals_vector` avec les catégories `"Elephant"`, `"Girafe"`, `"Singe"` et `"Cheval"`. Ici, il est impossible de dire qui est au-dessus ou au-dessous de l'autre.  

En revanche, les variables ordinales ont un ordre naturel. Considérons par exemple la variable catégorielle `temperature_vector` avec les catégories: `"Faible"`, `"Moyen"` et `"Fort"`. Ici, il est évident que `"Moyen"` se situe au-dessus de `"Faible"`, et  `"Fort"` se situe au-dessus de `"Moyen"`.



*** =instructions

Cliquez sur 'Submit Answer' pour observer comment R construit et affiche les variables nominales ou ordinales. Ne vous inquiétez pas si vous ne comprenez pas tout le code pour le moment.

*** =hint

Il suffit de cliquez sur le bouton 'Submit Answer' et de regarder la console. Observez comment R indique l'ordre des niveaux de facteurs pour les variables qualitatives ordonnées.


*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
animals_vector <- c("Elephant", "Girafe", "Singe", "Cheval")
temperature_vector <- c("Fort", "Faible", "Fort","Faible", "Moyen")

factor_animaux_vector <- factor(animals_vector)
factor_animaux_vector
factor_temperature_vector <- factor(temperature_vector, order = TRUE, levels = c("Faible", "Moyen", "Fort"))
factor_temperature_vector
```

*** =solution
```{r eval=FALSE}
animals_vector <- c("Elephant", "Girafe", "Singe", "Cheval")
temperature_vector <- c("Fort", "Faible", "Fort","Faible", "Moyen")

factor_animaux_vector <- factor(animals_vector)
factor_animaux_vector
factor_temperature_vector <- factor(temperature_vector, order = TRUE, levels = c("Faible", "Moyen", "Fort"))
factor_temperature_vector
```

*** =sct
```{r eval=FALSE}
msg = "Ne changez rien au code indiqu&#233; en exemple. Cliquez sur le bouton submit et regardez la solution ! "
test_object("animals_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("temperature_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_animaux_vector", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("factor_animaux_vector", incorrect_msg = msg)
test_object("factor_temperature_vector", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("factor_temperature_vector", incorrect_msg = msg)
success_msg("Vous avez compris ce qui s'est pass&#233; dans cet exercice ? G&#233;nial ! On passe &#224; l'exercice suivant pour regarder les facteurs et leur niveaux plus en d&#233;tails.")
```

*** =skills
1

--- type:NormalExercise xp:100 key:c8ac1a60ae
## Niveau de facteurs

Lorsque vous importez un jeu de données, vous remarquerez généralement qu'il contient des facteurs avec des niveaux spécifiques. Parfois vous allez vouloir changer les noms de ces niveaux, pour des raisons de clarté ou autre. R vous permet de le faire avec la fonction [`levels()`] (http://www.rdocumentation.org/packages/base/functions/levels) :


```
levels(factor_vector) <- c("nom1","nom2",...)
```
    



Un bon exemple est le fichier brut qui vous est fourni après une enquête terrain (un sondage par exemple). Une des questions standard pour chaque questionnaire est le sexe de la personne interrogée. Vous vous souvenez de la question précédente ? Il s'agit d'un facteur et ses niveaux sont souvent notés `"M"` et `"F"`. 


```
enquete_vector <- c("M","F","F","M","M")
```

Ensuite, lorsque vous voulez commencer votre analyse de données, votre principale préoccupation est de garder un bon aperçu de toutes les variables et de leur sens. À ce moment-là, vous voudrez sûrement changer les niveaux de facteur en indiquant `"Masculin"` et `"Feminin"` au lieu de `"M"` et `"F"` pour vous rendre la vie plus facile. 



*** =instructions

-Convertissez le vecteur de caractères `enquete_vector` dans un vecteur de facteurs. Assignez-le à « factor_enquete_vector ».
-Changez les niveaux de facteurs de `factor_enquete_vector` à `"Masculin"` et `"Feminin"`.


*** =hint
L'ordre dans lequel vous devez taper les niveaux de facteur est important. Conseil : regardez l'ordre dans lequel les niveaux sont affichés.
*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")

# Convertissez enquete_vector en facteur
factor_enquete_vector <- 

# Spécifiez les niveaux de  factor_enquete_vector
levels(factor_enquete_vector) <- 

factor_enquete_vector
```

*** =solution
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")

# Convertissez enquete_vector en facteurs
factor_enquete_vector <- factor(enquete_vector)

# Spécifiez les niveaux de  factor_enquete_vector
levels(factor_enquete_vector) <- c("Feminin", "Masculin")

factor_enquete_vector
```

*** =sct
```{r eval=FALSE}
msg = "Ne modifiez pas la d&#233;finition de <code>enquete_vector</code> ! "
test_object("enquete_vector", undefined_msg = msg, incorrect_msg = msg)
test_function("factor","x",
              incorrect_msg = "Avez-vous bien utilis&#233; la fonction <code>factor()</code> pour cr&#233;er <code>factor_enquete_vector</code> ? ")
test_object("factor_enquete_vector", eq_condition = "equal",
            undefined_msg = "Faites attention &#224; bien d&#233;finir la variable <code>factor_enquete_vector</code>.",
            incorrect_msg = paste("Avez-vous bien indiqu&#233; les bons niveaux de facteurs &#224; <code>factor_enquete_vector</code> ? ",
                                  "Souvenez-vous que R est sensible &#224; la diff&#233;rence entre les majuscules et les minuscules. Faites aussi attention &#224; respecter l'ordre des niveaux ! "))
success_msg("Magnifique ! On passe &#224; l'exercice suivant.")
```

*** =skills
1

--- type:NormalExercise xp:100 key:3b26500eac
## Décrire un factor

Après avoir terminé ce cours, une de vos fonctions préférées dans R sera [`summary()`] (http://www.rdocumentation.org/packages/base/functions/summary). Cela vous donnera un aperçu rapide de `une_variable` : 


```
summary(une_variable)
``` 

Pour en revenir à notre enquête, vous souhaitez savoir combien de réponses `"Masculin"` et combien de réponses `"Feminin"` vous avez dans votre étude. La fonction [`summary()`] (http://www.rdocumentation.org/packages/base/functions/summary) vous donne la réponse à cette question.

*** =instructions
Utilisez  [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary) sur les vecteurs `enquete_vector` et `factor_enquete_vector`. Interpréter les résultats de ces deux vecteurs. Sont-ils équivalents dans ce cas ? 

*** =hint
Utilisez  [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary) sur les vecteurs `enquete_vector` et `factor_enquete_vector`. C'est aussi simple que cela ! 

*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")
factor_enquete_vector <- factor(enquete_vector)
levels(factor_enquete_vector) <- c("Feminin", "Masculin")
factor_enquete_vector

# Placez votre code ici pour enquete_vector


# Placez votre code ici pour factor_enquete_vector

```

*** =solution
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")
factor_enquete_vector <- factor(enquete_vector)
levels(factor_enquete_vector) <- c("Feminin", "Masculin")
factor_enquete_vector

# Placez votre code ici pour enquete_vector
summary(enquete_vector)

# Placez votre code ici pour factor_enquete_vector
summary(factor_enquete_vector)
```

*** =sct
```{r eval=FALSE}
msg = "Ne changez rien aux 4 premi&#232;res lignes qui d&#233;finissent <code>enquete_vector</code> et <code>factor_enquete_vector</code>."
test_object("enquete_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_enquete_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_function("summary", "object", 
              incorrect_msg = c("Avez-vous utilis&#233; la fonction summary sur <code>enquete_vector</code> ? ",
                                "Avez-vous utilis&#233; la fonction summary sur <code>factor_enquete_vector</code> ? "))
success_msg("Bien ! Regardez la sortie dans la console. Le fait que vous ayez identifi&#233; <code>\"Masculin\"</code> et <code>\"Féminin\"</code> comme niveaux de facteur dans <code>factor_enquete_vector</code> permet &#224; R d'afficher le nombre d'&#233;l&#233;ments pour chaque cat&#233;gorie.")
```

*** =skills
1

--- type:NormalExercise xp:100 key:4fbc8f1e25
## Bataille des sexes

_"Tous les animaux sont égaux, mais certains animaux sont plus égaux que d'autres." -- George Orwell_

Dans `factor_enquete_vector`, nous avons un facteur à deux niveaux : Masculin et Féminin. Mais comment R compare t'il ces valeurs l'une par rapport à l'autre ? En d'autres termes, pour R qu'est-ce qui est le plus grand, Masculin ou Féminin ? 


*** =instructions
Lire le code dans l'éditeur, puis cliquez sur 'Submit Answer' pour voir si les hommes valent plus que les femmes.
*** =hint
Il suffit de cliquez sur 'Submit Answer' et jetez un oeil sur la sortie qui est affichée dans la console.
*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")
factor_enquete_vector <- factor(enquete_vector)
levels(factor_enquete_vector) <- c("Feminin", "Masculin")

# Masculin
factor_enquete_vector[1] 
# Féminin
factor_enquete_vector[2] 
# Guerre des sexes : Masculin est-il plus grand que Féminin ? 
factor_enquete_vector[1] > factor_enquete_vector[2] 
```

*** =solution
```{r eval=FALSE}
enquete_vector <- c("M", "F", "F", "M", "M")
factor_enquete_vector <- factor(enquete_vector)
levels(factor_enquete_vector) <- c("Feminin", "Masculin")

# Masculin
factor_enquete_vector[1] 
# Féminin
factor_enquete_vector[2] 
# Guerre des sexe : Masculin est il plus grand que Féminin ? 
factor_enquete_vector[1] > factor_enquete_vector[2] 
```

*** =sct
```{r eval=FALSE}
msg = "Ne changez rien aux 4 premi&#232;res lignes qui d&#233;finissent <code>enquete_vector</code> et <code>factor_enquete_vector</code>."
test_object("enquete_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_enquete_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("factor_enquete_vector[1] > factor_enquete_vector[2]", incorrect_msg = "Ne changez rien &#224; la derni&#232;re commande, nous voulons que cette bataille aie lieu une fois pour toutes ! ")
success_msg("Ouf, il semble que R soit neutre. Ou peut-&#234;tre veut-il juste rester en dehors de ces discussions ;-).")
```

*** =skills
1

--- type:NormalExercise xp:100 key:4cd29dda9c
## Les facteurs ordonnés

`"Masculin"` et `"Feminin"` étant des niveaux non ordonnés (ou nominaux) de facteurs, R retourne un message d'avertissement vous indiquant que le l'opérateur de comparaison n'est pas explicite. Comme vu auparavant, R attache une valeur égale aux niveaux de tels facteurs. 

Mais ce n'est pas toujours le cas ! Parfois il faudra aussi travailler avec des facteurs ordonnés, qui ont un ordre naturel dans les niveaux. Si c'est le cas, nous devons veiller à ce que R le comprenne... 

Disons que vous dirigez une équipe de recherche de cinq statisticiens et que vous souhaitez évaluer leurs performances. Pour ce faire, vous allez évaluer leur vitesse. Chaque statisticiens sera indiqué comme `"Lent"`, `"Rapide"` ou `"Ultra-Rapide"`. Vous stockerez les résultats dans `vitesse_vector`.


*** =instructions

Dans un premier temps, fabriquez `vitesse_vector`, sachant que : 

- Le Statisticien 1 est Rapide, 
- Le Statisticien 2 est Lent, 
- Le Statisticien 3 est Lent,
- Le Statisticien 4 est Rapide et
- Le Statisticien 5 est Ultra-Rapide.



Pas besoin pour l'instant de spécifier que ce sont des facteurs

*** =hint
Fabriquez `vitesse_vector`, il s'agit d'un vecteur contenant les chaînes des caractères : `"Rapide"`, `"Lent"`...


*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
# Fabriquez vitesse_vector
vitesse_vector <-
```

*** =solution
```{r eval=FALSE}
# Fabriquez vitesse_vector
vitesse_vector <- c("Rapide", "Lent", "Lent", "Rapide", "Ultra-Rapide")
```

*** =sct
```{r eval=FALSE}
test_object("vitesse_vector", 
            undefined_msg = "Faites attention &#224; bien d&#233;finir la variable <code>vitesse_vector</code>.",
            incorrect_msg = "Faites attention &#224; bien assigner le bon vecteur &#224; vitesse_vector. Souvenez-vous que R est sensible &#224; la différence entre les majuscules et les minuscules ! ")
success_msg("Bien jou&#233; ! On passe &#224; l'exercice suivant.")
```

*** =skills
1

--- type:NormalExercise xp:100 key:7b2bf68c54
## Les facteurs ordonnés (2)

`vitesse_vector` doit être converti en un facteur ordinal puisque ses niveaux ont un ordre naturel. Par défaut, la fonction [`factor()`] (http://www.rdocumentation.org/packages/base/functions/factor) transforme `vitesse_vector` en un factor non ordonné. Pour créer un factor ordonné, vous devez ajouter deux arguments supplémentaires: `ordered` et `levels`. 




```
factor(un_vector, ordered = TRUE, levels = c("Level_1", "Level_2"...))
``` 


En définissant l'argument `ordered` à `TRUE` dans la fonction [`factor()`] (http://www.rdocumentation.org/packages/base/functions/factor), vous indiquez que le factor est ordonné. Avec l'argument `levels` vous donnez les valeurs du factor dans l'ordre correct, par exemple `levels = c("Faible", "Moyen", "Fort")`.



*** =instructions

L'exemple de code ci-dessus permet de définir une variable `factor_vitesse_vector` en tenant compte du fait qu'il y a un ordre spécifique pour les niveaux de facteur. Ne changez rien aux autres lignes de code, cela imprimera le facteur dans la console et générera un résumé de celui-ci.



*** =hint

Utilisez la fonction [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor) pour fabriquer `factor_vitesse_vector` à partir de `speed_character_vector`. l'argument order doit être à `TRUE` puisqu'il y a un ordre. Les niveaux de facteurs sont alors: `c("Lent", "Rapide", "Ultra-Rapide")`.




*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
vitesse_vector <- c("Rapide", "Lent", "Lent", "Rapide", "Ultra-Rapide") 

# Rajoutez votre code ci dessous
factor_vitesse_vector <-  

# Affichez
factor_vitesse_vector

# R affiche automatiquement le bon ordre
summary(factor_vitesse_vector) 
```

*** =solution
```{r eval=FALSE}
vitesse_vector <- c("Rapide", "Lent", "Lent", "Rapide", "Ultra-Rapide")

# Rajoutez votre code ci dessous
factor_vitesse_vector <- factor(vitesse_vector, ordered = TRUE, levels= c("Lent", "Rapide", "Ultra-Rapide"))

# Visualisation
factor_vitesse_vector

# R affiche automatiquement le bon ordre
summary(factor_vitesse_vector)
```

*** =sct
```{r eval=FALSE}
msg = "Ne changez rien au code qui d&#233;fini <code>vitesse_vector</code>."
test_object("vitesse_vector", undefined_msg = msg, incorrect_msg = msg)
test_correct({
  test_object("factor_vitesse_vector", eq_condition = "equal",
              undefined_msg = "Faites attention &#224; bien d&#233;finir la variable <code>factor_vitesse_vector</code>.",
              incorrect_msg = "Faites attention &#224; bien assigner les bons niveaux &#224; factor_vitesse_vector. Faites aussi attention &#224; l'ordre des &#233;l&#233;ments dans le param&#224;tre levels.")
},{
  test_function("factor", c("x", "levels", "ordered"))
})
success_msg("Bien ! Regardez dans la console. On voit maintenant qu'il y a une notion d'ordre entre les niveaux de facteurs, on voit cela gr&#226;ce au signe <code><</code>. On passe &#224; l'exercice suivant.");
```

*** =skills
1

--- type:NormalExercise xp:100 key:09dff7e67b
## Comparaison de facteurs ordonnés


Le statisticien numéro 2 entre dans votre bureau et commence à se plaindre que le statisticien numéro 5 ralentit l'ensemble du projet. Puisque vous savez que le statisticien numéro 2 a la réputation d'être un malin, vous décidez tout d'abord de vérifier si sa déclaration est correcte. 

Le fait que `factor_vitesse_vector` est maintenant ordonné permet de comparer les différents éléments (les statisticiens dans ce cas). Vous pouvez faire assez simplement cela en utilisant les opérateurs connus.



*** =instructions

Vérifiez si le statisticien numéro 2 est plus rapide que statisticien numéro 5 et assignez le résultat à `la_comparaison`. N'oubliez pas l'opérateur `>` permet de vérifier si un élément est plus grand que l'autre.

*** =hint
Vous pouvez comparer les éléments du vecteur avec l'opérateur `>`. 

```
vector[1] > vector[2]
``` 
Vérifie si le premier élément du vecteur est plus grand que le deuxième élément.

*** =pre_exercise_code
```{r eval=FALSE}
# no pec
```

*** =sample_code
```{r eval=FALSE}
vitesse_vector <- c("Rapide", "Lent", "Lent", "Rapide", "Ultra-Rapide")
factor_vitesse_vector <- factor(vitesse_vector, ordered = TRUE, levels = c("Lent", "Rapide", "Ultra-Rapide"))

# Votre code
la_comparaison <- 

# Est-ce que le statisticien 2 est plus rapide que le 5 ? 
la_comparaison
```

*** =solution
```{r eval=FALSE}
vitesse_vector <- c("Rapide", "Lent", "Lent", "Rapide", "Ultra-Rapide")
factor_vitesse_vector <- factor(vitesse_vector, ordered = TRUE, levels = c("Lent", "Rapide", "Ultra-Rapide"))

# Votre code
la_comparaison <- factor_vitesse_vector[2] > factor_vitesse_vector[5] 

# Est-ce que le statisticien 2 est plus rapide que le 5 ? 
la_comparaison
```

*** =sct
```{r eval=FALSE}
msg = "Ne changez rien aux instructions qui d&#233;finissent <code>vitesse_vector</code> et <code>factor_vitesse_vector</code> ! "
test_object("vitesse_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_vitesse_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("la_comparaison",
            undefined_msg = "Faites attention &#224; bien d&#233;finir la variable <code>la_comparaison</code>.",
            incorrect_msg = "Il semblerait que <code>la_comparaison</code> ne soit pas correctement effectu&#233;e. V&#233;rifiez les indices.");
success_msg("G&#233;nial ! Que nous apprend ce r&#233;sultat ? le statisticien 2 se plaint au sujet du num&#233;ro 5 alors qu'en fait il est celui qui fait tout ralentir ! Voil&#224; qui termine le chapitre sur les facteurs. Avec une base solide sur les vecteurs, les matrices et les facteurs, vous &#234;tes pr&#234;ts &#224; plonger dans le monde magnifique des data.frame, une structuration des donn&#233;es tr&#232;s importante en R ! ")
```

*** =skills
1
